// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    println!("Hello there!")
}

// Pour l'intro1.rs l'erreur dans le code est l'absence de point-virgule à la fin du dernier appel à "println!" dans la fonction main.
//En Rust, chaque instruction doit se terminer par un point-virgule sauf si elle est suivie par un bloc de code ou
//si elle est la dernière expression d'une fonction renvoyant une valeur

//Pour l'intro2.rs , l'erreur dans le code se trouve dans l'utilisation incorrecte de la macro printline!. 
//En Rust, il n'existe pas de macro appelée printline!. La macro correcte pour imprimer une ligne de texte à la console est println!
